***Important*** DO NOT CLEAR THE OUTPUT OF THIS NOTEBOOK AFTER EXECUTION!!!

```python
# if the following command generates an error, you probably didn't enable
# the cluster security option "Allow API access to all Google Cloud services"
# under Manage Security → Project Access when setting up the cluster
!gcloud dataproc clusters list --region us-central1
```

Output:
```
NAME          PLATFORM  PRIMARY_WORKER_COUNT  SECONDARY_WORKER_COUNT  STATUS   ZONE           SCHEDULED_DELETE
cluster-35c2  GCE       2                                             RUNNING  us-central1-a
```

# Imports & Setup

```python
!pip install -q google-cloud-storage==1.43.0
!pip install -q graphframes
```

Output:
```
[33mWARNING: Running pip as the 'root' user can result in broken permissions and conflicting behaviour with the system package manager. It is recommended to use a virtual environment instead: https://pip.pypa.io/warnings/venv[0m[33m
[0m[33mWARNING: Running pip as the 'root' user can result in broken permissions and conflicting behaviour with the system package manager. It is recommended to use a virtual environment instead: https://pip.pypa.io/warnings/venv[0m[33m
[0m
```

```python
import pyspark
import sys
from collections import Counter, OrderedDict, defaultdict
import itertools
from itertools import islice, count, groupby
import pandas as pd
import os
import re
from operator import itemgetter
import nltk
from nltk.stem.porter import *
from nltk.corpus import stopwords
from time import time
from pathlib import Path
import pickle
import pandas as pd
from google.cloud import storage

import hashlib
def _hash(s):
    return hashlib.blake2b(bytes(s, encoding='utf8'), digest_size=5).hexdigest()

nltk.download('stopwords')
```

Output:
```
[nltk_data] Downloading package stopwords to /root/nltk_data...
[nltk_data]   Package stopwords is already up-to-date!
True
```

```python
# if nothing prints here you forgot to include the initialization script when starting the cluster
!ls -l /usr/lib/spark/jars/graph*
```

Output:
```
-rw-r--r-- 1 root root 247882 Feb 20 09:06 /usr/lib/spark/jars/graphframes-0.8.2-spark3.1-s_2.12.jar
```

```python
from pyspark.sql import *
from pyspark.sql.functions import *
from pyspark import SparkContext, SparkConf, SparkFiles
from pyspark.sql import SQLContext
from graphframes import *
```

```python
spark
```

Output:
```
<pyspark.sql.session.SparkSession at 0x7f3220f9d780>
```

<!-- grade_id: cell-bucket_name, points: 0, grade: false, solution: true -->
```python
# Put your bucket name below and make sure you can access it without an error
bucket_name = 'bucket_title'
full_path = f"gs://{bucket_name}/"
paths=[]

client = storage.Client()
blobs = client.list_blobs(bucket_name)
for b in blobs:
    if b.name != 'startup_script_gcp.sh':
        paths.append(full_path+b.name)
```

***GCP setup is complete!*** If you got here without any errors you've earned 10 out of the 35 points of this part.

# Building an inverted index

Here, we read the entire corpus to an rdd, directly from Google Storage Bucket and use your code from Colab to construct an inverted index.

```python
parquetFile = spark.read.parquet(*paths)
doc_text_pairs = parquetFile.select("title", "id").rdd
```

We will count the number of pages to make sure we are looking at the entire corpus. The number of pages should be more than 6M

```python
# Count number of wiki pages
parquetFile.count()
```

Output:
```
6348910
```

Let's import the inverted index module. Note that you need to use the staff-provided version called `inverted_index_gcp.py`, which contains helper functions to writing and reading the posting files similar to the Colab version, but with writing done to a Google Cloud Storage bucket.

```python
# if nothing prints here you forgot to upload the file inverted_index_gcp.py to the home dir
%cd -q /home/dataproc
!ls inverted_index_gcp.py
```

Output:
```
inverted_index_gcp.py
```

```python
# adding our python module to the cluster
sc.addFile("/home/dataproc/inverted_index_gcp.py")
sys.path.insert(0,SparkFiles.getRootDirectory())
```

```python
from inverted_index_gcp import InvertedIndex
```

**YOUR TASK (10 POINTS)**: Use your implementation of `word_count`, `reduce_word_counts`, `calculate_df`, and `partition_postings_and_write` functions from Colab to build an inverted index for all of English Wikipedia in under 2 hours.

A few notes:
1. The number of corpus stopwords below is a bit bigger than the colab version since we are working on the whole corpus and not just on one file.
2. You need to slightly modify your implementation of  `partition_postings_and_write` because the signature of `InvertedIndex.write_a_posting_list` has changed and now includes an additional argument called `bucket_name` for the target bucket. See the module for more details.
3. You are not allowed to change any of the code not coming from Colab.

<!-- grade_id: cell-token2bucket, points: 0, grade: false, solution: true -->
```python
english_stopwords = frozenset(stopwords.words('english'))
corpus_stopwords = ["category", "references", "also", "external", "links",
                    "may", "first", "see", "history", "people", "one", "two",
                    "part", "thumb", "including", "second", "following",
                    "many", "however", "would", "became"]

all_stopwords = english_stopwords.union(corpus_stopwords)
RE_WORD = re.compile(r"""[\#\@\w](['\-]?\w){2,24}""", re.UNICODE)

NUM_BUCKETS = 124
def token2bucket_id(token):
  return int(_hash(token),16) % NUM_BUCKETS

def word_count(text, id):
  tokens = [token.group() for token in RE_WORD.finditer(text.lower())]
  # Filter out stopwords
  tokens = [token for token in tokens if token not in all_stopwords]
  tf = Counter(tokens)
  # Convert the Counter object to a list of (token, (doc_id, tf)) tuples
  return [(token, (id, freq)) for token, freq in tf.items()]

def calculate_df(postings):
  df_rdd = postings.map(lambda x: (x[0], len(x[1])))
  return df_rdd

def reduce_word_counts(unsorted_pl):
  sorted_pl = sorted(unsorted_pl, key=lambda x: x[0])
  return sorted_pl

def partition_postings_and_write(postings, bucket_name=bucket_name):
    postings_to_buckets = postings.map(lambda x: (token2bucket_id(x[0]), (x[0], x[1])))
    postings_by_bucket = postings_to_buckets.groupByKey().mapValues(list)
    posting_locations = postings_by_bucket.map(
        lambda bucket: InvertedIndex.write_a_posting_list((bucket[0], list(bucket[1])), bucket_name)
    )
    return posting_locations
```

<!-- grade_id: cell-index_construction, points: 0, grade: false, solution: true -->
```python
# time the index creation time
t_start = time()
# word counts map
word_counts = doc_text_pairs.flatMap(lambda x: word_count(x[0], x[1]))
postings = word_counts.groupByKey().mapValues(reduce_word_counts)
# filtering postings and calculate df
postings_filtered = postings.filter(lambda x: len(x[1])>50)
w2df = calculate_df(postings_filtered)
w2df_dict = w2df.collectAsMap()
# partition posting lists and write out
_ = partition_postings_and_write(postings_filtered).collect()
index_const_time = time() - t_start
```

<!-- grade_id: cell-index_const_time, points: 10, grade: true, solution: false -->
```python
# test index construction time
assert index_const_time < 60*120
```

<!-- grade_id: collect-posting, points: 0, grade: true, solution: false -->
```python
# collect all posting lists locations into one super-set
super_posting_locs = defaultdict(list)
for blob in client.list_blobs(bucket_name, prefix='postings_gcp_title'):
  if not blob.name.endswith("pickle"):
    continue
  with blob.open("rb") as f:
    posting_locs = pickle.load(f)
    for k, v in posting_locs.items():
      super_posting_locs[k].extend(v)
```

Putting it all together

```python
# Create inverted index instance
inverted = InvertedIndex()
# Adding the posting locations dictionary to the inverted index
inverted.posting_locs = super_posting_locs
# Add the token - df dictionary to the inverted index
inverted.df = w2df_dict
# write the global stats out
inverted.write_index('.', 'index')
# upload to gs
index_src = "index.pkl"
index_dst = f'gs://{bucket_name}/postings_gcp_title/{index_src}'
!gsutil cp $index_src $index_dst
```

Output:
```
Copying file://index.pkl [Content-Type=application/octet-stream]...
/ [1 files][ 18.4 MiB/ 18.4 MiB]                                                
Operation completed over 1 objects/18.4 MiB.
```

```python
!gsutil ls -lh $index_dst
```

Output:
```
18.45 MiB  2024-02-20T12:05:42Z  gs://ex3adar/postings_gcp/index.pkl
TOTAL: 1 objects, 19343013 bytes (18.45 MiB)
```

# PageRank

**YOUR TASK (10 POINTS):** Compute PageRank for the entire English Wikipedia. Use your implementation for `generate_graph` function from Colab below.

```python
# Put your `generate_graph` function here
def generate_graph(pages):
  ''' Compute the directed graph generated by wiki links.
  Parameters:
  -----------
    pages: RDD
      An RDD where each row consists of one wikipedia articles with 'id' and
      'anchor_text'.
  Returns:
  --------
    edges: RDD
      An RDD where each row represents an edge in the directed graph created by
      the wikipedia links. The first entry should the source page id and the
      second entry is the destination page id. No duplicates should be present.
    vertices: RDD
      An RDD where each row represents a vetrix (node) in the directed graph
      created by the wikipedia links. No duplicates should be present.
  '''
  # Initialize vertices
  sources = pages.map(lambda x:(x[0], ))
  dests = pages.flatMap(lambda x:x[1]).map(lambda x:(x[0], ))
  vertices = sources.union(dests).distinct()
  # Construct edges
  edges = pages.flatMapValues(lambda x:x)
  edges = edges.map(lambda x:(x[0], x[1][0])).distinct()
  return edges, vertices
```

<!-- grade_id: cell-PageRank, points: 0, grade: false, solution: true -->
```python
t_start = time()
#pages_links = spark.read.parquet("gs://wikidata_preprocessed/*").select("id", "anchor_text").rdd
pages_links = parquetFile.select ("id","anchor_text").rdd
# construct the graph
edges, vertices = generate_graph(pages_links)
# compute PageRank
edgesDF = edges.toDF(['src', 'dst']).repartition(124, 'src')
verticesDF = vertices.toDF(['id']).repartition(124, 'id')
g = GraphFrame(verticesDF, edgesDF)
pr_results = g.pageRank(resetProbability=0.15, maxIter=6)
pr = pr_results.vertices.select("id", "pagerank")
pr = pr.sort(col('pagerank').desc())
pr.repartition(1).write.csv(f'gs://{bucket_name}/pr', compression="gzip")
pr_time = time() - t_start
pr.show()
```

Output:
```
24/02/20 12:34:13 WARN YarnAllocator: Container from a bad node: container_1708419947277_0002_01_000019 on host: cluster-35c2-w-1.c.wired-climate-414717.internal. Exit status: 143. Diagnostics: [2024-02-20 12:34:13.326]Container killed on request. Exit code is 143
[2024-02-20 12:34:13.327]Container exited with a non-zero exit code 143. 
[2024-02-20 12:34:13.328]Killed by external signal
.
24/02/20 12:34:13 ERROR YarnScheduler: Lost executor 19 on cluster-35c2-w-1.c.wired-climate-414717.internal: Container from a bad node: container_1708419947277_0002_01_000019 on host: cluster-35c2-w-1.c.wired-climate-414717.internal. Exit status: 143. Diagnostics: [2024-02-20 12:34:13.326]Container killed on request. Exit code is 143
[2024-02-20 12:34:13.327]Container exited with a non-zero exit code 143. 
[2024-02-20 12:34:13.328]Killed by external signal
.
24/02/20 12:34:13 WARN TaskSetManager: Lost task 141.0 in stage 20.0 (TID 1614) (cluster-35c2-w-1.c.wired-climate-414717.internal executor 19): E
```

<!-- grade_id: cell-PageRank_time, points: 10, grade: true, solution: false -->
```python
# test that PageRank computaion took less than 1 hour
assert pr_time < 120*60
```

# Reporting

**YOUR TASK (5 points):** execute and complete the following lines to complete
the reporting requirements for assignment #3.

```python
# size of input data
!gsutil du -sh "gs://wikidata_preprocessed/"
```

Output:
```
14.28 GiB    gs://wikidata_preprocessed
```

```python
# size of index data
index_dst = f'gs://{bucket_name}/postings_gcp_title/'
!gsutil du -sh "$index_dst"
```

Output:
```
5.92 GiB     gs://ex3adar/postings_gcp
```

```python
# How many USD credits did you use in GCP during the course of this assignment?
cost = 0
print(f'I used {cost} USD credit during the course of this assignment')
```

Output:
```
I used 0 USD credit during the course of this assignment
```

**Bonus (10 points)** if you implement PageRank in pure PySpark, i.e. without using the GraphFrames package, AND manage to complete 10 iterations of your algorithm on the entire English Wikipedia in less than an hour.


<!-- grade_id: cell-PageRank_Bonus, points: 0, grade: false, solution: true -->
```python
#If you have decided to do the bonus task - please copy the code here

bonus_flag = False # Turn flag on (True) if you have implemented this part

t_start = time()

# PLACE YOUR CODE HERE

pr_time_Bonus = time() - t_start

```

<!-- grade_id: cell-PageRank_Bonus-time, points: 10, grade: true, solution: false -->
```python
# Note:test that PageRank computaion took less than 1 hour
assert pr_time_Bonus < 60*60 and bonus_flag
```

Output:
```
AssertionError                            Traceback (most recent call last)
/tmp/ipykernel_10293/574311967.py in <cell line: 2>()
      1 # Note:test that PageRank computaion took less than 1 hour
----> 2 assert pr_time_Bonus < 60*60 and bonus_flag

AssertionError:
```

```python

```